<problem>
  <text>
    <script type="loncapa/python">
# variable names and values
# if multiple variables, the number of values to test should be the same
variable_values = {'n':[1,2,3], 'k':[4,5,6]}


# random variables
r = random.randrange(3,6,1)

# Solutions with variables converted to string
# Make sure you name the solution with part id at the end. e.g. 'solution1' will be solution for part 1.
solution1 = "(26^n)-(25^n)+{0}".format(r)
solution2 = "8"
solution3 = "35*k"

# Group all solutions into a list
solutions = [solution1, solution2, solution3]



    </script>


<ul>
<li>How many strings of length $n consisting of lower case English letters, and including at least one "x", are there?</li>
</ul>
    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
    </customresponse>

<ul>
<li>Another part, the answer is?</li>
</ul>
    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
    </customresponse>

<hr>
<ul>
<li>Part 3, fill in your answer here.</li>
</ul>
    <customresponse cfn="check" expect="\[$solution3\]">
      <textline/>
    </customresponse>


  </text>
</problem>
